---
title: Test
description: Test
---

## Sample exercise

```yaml
type: NormalExercise
key: e7d96c5b57
lang: python
xp: 100
skills: 2
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: a4596c8ecd
xp: 100
```

In this first and preliminary attempt, the student is asked to know how to declare variables

`@instructions`
## How to code variables, what are alternatives in this case of situation

`@hint`


`@pre_exercise_code`
```{Python}

```

`@sample_code`
```{Python}

```

`@solution`
```{Python}

```

`@sct`
```{Python}

```
